---
title       : Which movies categories gets better movie ratings? Creating with Categorical and Dummy Variables for Decision Trees
description : Introduces categorical variable and methods to convert movies dataset for Decision Tree

--- type:MultipleChoiceExercise lang:python xp:50 skills:2
## Which variables are categorical by name?

Introduction to categorical variable concepts and why for logistic regression? - load dataset columns head()

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint

*** =pre_exercise_code
```{python}
# pec
```

*** =sct
```{python}
test_mc(2) # if 2 is the correct option.
```

--- type:NormalExercise lang:python xp:100 skills:2
## Methods to check categorical by value

Use panda functions to determine which variables are categorical - describe() and unique() and plot to check

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
success_msg("Great work!")
```

--- type:NormalExercise lang:python xp:100 skills:2
## Create panda series categorical values; pandas 'categorical' var type

use pandas categorical method to create a series of categorical variables

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
success_msg("Great work!")
```

--- type:NormalExercise lang:python xp:100 skills:2
## Convert variables into categorical values using Get_Dummy

Use pandas get_dummy method to convert any one variables and join back to dataframe

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
success_msg("Great work!")
```

--- type:NormalExercise lang:python xp:100 skills:2
## Use dataframe with created categorical variables to produce a decision tree

Use decision tree with dummy values to determine what categories produce higher ratings

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
success_msg("Great work!")
```

